```python
# Initialize Otter
import otter
grader = otter.Notebook("controlFunctions.ipynb")
```

## Lecture Section

### Control

Control statements in Python manage how the code executes. It allows decision-making, repeating tasks, and handling errors.

For the most part, this course assumes that you are familiar with at least one other programming language. Therefore, much of this lecture may be redundant. It is a longer lecture, but much of it should be review. I trust that you can glean the information you need from it!

This lecture will cover control aspects of Python, and functions. We will cover:
* Control statements, including:
    * if/elif/else-statements
        * logical operators
    * for-loops
    * while-loops
* `continue` & `break`
* Functions, including:
    * Defining and usage
    * Arguments and parameters
    * `return`
    * Global & local variables
    * Mutability with functions

#### Conditional Statements

The simplest control statement are made of a sequence of `if`, `elif`, and `else`. These make conditional statements. Next to `if` and `elif` should be a condition, allowing the code indented beneath to run only if the condition is met (the condition returns `True`).

**Note on indentation**: Python does not use brackets to manage control or code-blocks. Instead, it uses indents. Most IDEs will automatically indent when necessary, but you can indent by pressing your `tab` key.

Let's try an example with only `if`. Change the `x` variable to `5` and see what happens!

```python
x = 10
if x == 5:
    print('x is a 5!')
```

If we want to check conditions sequentially, and we want to only run the first `True` condition, we can add `elif` statements.

```python
x = 10
if x == 5:
    print('x is a 5!')
elif x == 10:
    print('x is a 10!')
```

Let's try something a little more complicated:

```python
x = 10
if x %2 == 0:
    print('x is divisible by 2!')
elif x == 10:
    print('x is a 10!')
```

Notice how the if-conditional is the only code-block that runs. This is because it is the first one to return `True` among the `if` and `elif` statements belonging to the same control sequence and indentation level.

The next example has two control sequences, so both will print.

```python
x = 10
if x %2 == 0:
    print('x is divisible by 2!')
# new control sequence!
if x == 10:
    print('x is a 10!')
```

We can also nest `if/elif/else` sequences

```python
x = 20
if x %2 == 0:
    if x == 10:
        print('x is a 10!')
    elif x == 20:
        print('x is a 20!')
```

What about `else`? Great question!

`else` always comes at the end of a an `if/elif/else` control sequence, and runs if not of the conditions in the `if` or `elif` statements return `True`


```python
x = 30
if x %2 == 0:
    if x == 10:
        print('x is a 10!')
    elif x == 20:
        print('x is a 20!')
    else:
        print('x is not a 10 or a 20 but is divisible by 2!')
```

```python
x = 35
if x % 2 == 0:
    if x == 10:
        print('x is a 10!')
    elif x == 20:
        print('x is a 20!')
    else:
        print('x is not a 10 or a 20 but is divisible by 2!')
else:
    print('x is not divisible by 2!')
```

For conditional sequences, you always need **only 1** `if`, **0 or more** `elif`, and **at most 1** `else`.

```python
#### Logic
```

Let's quickly talk about logic.

When using if/elif/else or `=`, we can string multiple logical sequences together using `and` and `or`. If we want something negated, we can use `not` or `!` (when pairing with `!=`).

* `and` means the statements on both sides must be true
* `or` means one or both of the statements on either side need to be true
* `not` or `!` negates the truth of the statement

```python
l = [1, 2, 3, 4, 5]
for item in l:
    if item % 2 == 0 and item != 4:
        print(item)
```

```python
for item in l:
    if item % 2 == 0 and not item == 4:
        print(item)
```

Use parentheses to collect logic flow. In the example below, either `item == 4` or `item != 1 and item != 2` needs to be true. The number `4` satisfies the first, and `3` and `5` satisfy the second because those numbers are not 1 **and** they are not 2 (both sides must be true).

```python
for item in l:
    if item == 4 or (item != 1 and item != 2):
        print(item)
```

If we swap the `and` to an `or`, we get a very different result. Now, either `item == 4` or `item != 1 or item != 2` needs to be true in order to print. Similarly, only one side of `item != 1 or item != 2` needs to be true for the whole statement to be true! Since `1 != 2`, it is printed. `2 != 1`, `3` does not equal either, and neither does `5`. `4` prints because of `item == 4`. Be very careful when you are creating logic chains!

```python
for item in l:
    if item == 4 or (item != 1 or item != 2):
        print(item)
```

The last not on logic is a quick helper. In Python, when you have an `or` statement, if the first half is true, the second isn't read. Similarly, with `and` statements, if the first is false, the second isn't read. This is helpful if you are doing an operation that cannot occur on a certain type of variable.


Try it out below!

```python
x = ['hi', 1]
for item in x:
    if type(item) == int and item % 1 == 0:
        print(item)
```

```python
x = ['hi', 1]
for item in x:
    if  item % 1 == 0 and type(item) == int :
        print(item)
```

```python
x = ['hi', 1]
for item in x:
    if  item == 'hi' or item % 1 == 0 :
        print(item)
```

```python
x = ['hi', 1]
for item in x:
    if  item % 1 == 0 or item == 'hi' :
        print(item)
```

#### While-loops

Similar to `if` statements, `while` allows the code indented beneath to happen as long as the condition is `True`. However, the code beneath will repeat until the condition is `False`.

Normally, we increment a variable a certain number of times. This variable is usually `i`, but it doesn't have to be!

Be careful with the code below. You can make it loop infinitely by changing the condition to `while True`, but it may crash your assignment. Your auto-grader and Gradescope have time-out periods that will trigger if you leave an infinite loop, so change it back if you do it!

```python
i = 0
while i < 10:
    print(i)
    i += 1
```

#### For-loops

`For`-loops are similar to `while` loops because they repeat the indented coded beneath them, too. However, `for`-loops are used for iterating over a sequence.

This prevents them from infinite looping, because sequences are finite. However, they don't automatically stop when a condition changes, and you can't change the sequence while it's being iterated over, or you will get unpredictable or undefined behavior.

This type of behavior means... it's unlikely to cause a formal error, but it is likely to work different than you expected!

Try the code below, but try changing the string to a list!

```python
x = "a sequence"
for i in x:
    print(i)
```

If we want the `for` loop to iterate over a numerical range, we can use `range`. It starts at 0, and goes to the number given but does not include the number given.

```python
for i in range(10):
    print(i)
```

Both `while` and `for` loops offer further control with `continue`, `break`, and `else`.

1. `continue` stops the code execution and goes back to the top of the `while` loop, or increments and moves ahead in the `for` loop's sequence.
2. `end` completely stops the loop and moved on to the code in the rest of the file.
3. `else` will run the code indented below it after a `while` or `for` loop has naturally finished. Naturally means without error, and without `end`.

```python
x = "a s equ en ce"
for i in x:
    if i == " ":
        continue
    print(i)
```

```python
x = "a sequence"
for i in x:
    if i == " ":
        break
    print(i)
```

```python
x = "a sequence"
for i in x:
    print(i)
else:
    print("all done!")
```

```python
while True:
    break
else:
    print("i'm not going to work, the loop ended `unaturally`")
```

There are other aspects of control that we will cover in later lessons, but combining these should give you a good start.

If this is your first programming language, I highly recommend doing some of the online exercises or tutorials at the end of this lesson. A lot of programming is critical thinking and problem-solving with nested control sequences. Practice, practice, practice!

### Functions


A function lets you module your code. The code indented under a function only runs whe you call the function.

To create a function, we use the `def` keyword. Next, we type the name of our function, and `():`. Inside our parentheses, we can place any number of `parameters`, which allow us to access `global` variables that were created after the function was made, that may change after the function is defined, or `local` variables that exist elsewhere. The values we supply the `parameters` with are called `arguments`.

Functions can access `global` variables defined before it, `local` and `global` variables that are passed as arguments to the parameters, and `local` variables that are defined within it.
* **Global**: defined outside of functions, control statements (usually), and can be accessed by any program in the script/file.
*  **Local**: can only be used within the function/class and is not accessible anywhere else

To run the function code, we call the name of the function with the arguments we want to pass in. If the parameter doesn't have a `default` value, we are required to give it an argument.

```python
global_variable = "global variable"
def printer(): #should never do it this way, but it will work
    print(global_variable)

printer()
```

Pay attention to the code below. This is your usual function use-case. You define a function, then you call it some time later. The argument does not need to have the same name as the variable you give it- that would be very limiting and counter-intuitive! In our example, the argument is `x`, and we pass the variable `global_variable2` to it.

Note that `global_variable2` was created after the function. The function can still access it because we pass it as an argument, `x`. In the function, we use 'x' where we want our variable. `x` represents everything we might pass. If you uncomment the other function call, `x` becomes `global_variable3` on its second call!



```python
def printer(x): #make an argument instead
    print(x)

global_variable2 = "global variable2"
printer(global_variable2)
global_variable3 = "global variable3"
#printer(global_variable3)
```

The `printer()` function below uses a `default` argument. If an argument isn't passed, it assumes the `default` value.

```python
def printer(x='nothing was passed to me :('):
    print(x)

printer()
printer("hi")
```

We can have any number of arguments and types, too.

```python
def printer(i, x='repeat me!'):
    for n in range(i): # range lets
        print(x)

printer(5)
```

Parameters that have `default` values need to be defined first. You can place your arguments in any order when you call the function, though, as long as you define which parameter applies to which variables.

```python
def printer(i, x='repeat me!'):
    for n in range(i): # range lets
        print(x)

printer(x="i'm first now", i=5)
```

You can call functions from other functions, too!

```python
def print_first():
    print('first')

def printer():
    print_first()
    print("second")

printer()
```

You can also nest functions, but it's generally not good practice because it's easy to lose locality and access, so I won't cover it here.

#### `return`

Finally, functions can return a variable after it's called. If we don't define a `return` value, it automatically returns `None`. In the examples above, we did not have a `return` statement.

```python
def printer(x): #returns None
    print(x)

y = printer('hi')
print(y)
```

**Note**: in a notebook format, `None` won't show unless you use `print()`

```python
def printer(x): #returns x
    print(x)
    return x

y = printer('hi')
print(y)
```

```python
def add_3_numbers(x, y, z): #another example
    return x + y + z

total = add_3_numbers(10, 20, 30)
total
```

Finally, we can return more than one thing. We can collect these items in more than one way, too.

```python
def add_multiply_3_numbers(x, y, z): #another example
    return x + y + z, x*y*z

total = add_multiply_3_numbers(10, 20, 30)
total
```

```python
def add_multiply_3_numbers(x, y, z): #another example
    return x + y + z, x*y*z

sum, mult = add_multiply_3_numbers(10, 20, 30)
print(sum)
print(mult)
```

#### Mutability & Functions

Mutable data structures, like lists, are a little complicated with functions. Since they are mutable, we don't need to return them in order to change them. Our paremter `l` becomes our argument `my_l` and changes it `in-place`.

```python
def mutate_l(l):
    l[1] = 7

my_l = [1,2,3]
mutate_l(my_l)
my_l
```

This isn't possible with immutable variable, like strings.

```python
def mutate_s(s):
    s = s[0:1]

my_s = 'hello'
mutate_l(my_s)
my_s
```

Instead, we need to return and redefine our variable.

```python
def mutate_s(s):
    s = s[0:1]
    return s

my_s = 'hello'
my_s = mutate_s(my_s)
print(my_s)
```

## Assignments Section

**Question 1.** You will create a function called `return_div_3`. It should have one required (no default) parameter that takes a list. The function should iterate through the list and find the items divisible by 3. It should place those items into a new list. The new list should be returned. For example, if you pass `[1, 2, 3]` to the function, it should return `[3]`. The list can be made of mixed types, so make sure you ignore anything that isn't an integer.

```python
def return_div_3(l):
    pass

```

```python
grader.check("q1")
```

**Question 2.**

**Question 2.** For this question, you will create a function called `extract_encouragement()`.

* It should have one parameter, which takes a list.
* It will `return` a string

There is a list of words in a variable called `words`. Within this list is a secret message. The message can be created by keeping the words at the prime-numbered indexes (2, 3, 5, 7...)

You must change the original list in-place. Do not return it! Instead, you will return the hidden message as a string. Each word should be separated with a space. Do not change anything else!

Notes:
* ''.join(string) is an easy way to turn a list of words into a string
* You may use a helper function from a library to determine if a number is prime, or you can create your own
* You may find `.sort()` (in-place) or `sorted()` (must be assigned a new variable) helpful
* Lists change in-place when you pass them to a function, so you do not need to return it. The `words` variable is there so you have an example - the test cases use other lists.
* Hardcoding is absolutely not allowed. Your points will be removed if you do not code the solution and/or if you manually count the indexes to return the string. The code should work for any hidden message where the words exist at any numbers of prime indexes

```python
words = [
    "banana", "table", "you", "are", "cloud", "doing", "river", "great", "apple", "work",
    "keep", "on", "shining", "this", "and", "trust", "in", "assignment;", "path", "keep", "always",
    "believe", "mirror", "working", "can", "prosper", "anything", "with", "kindness", "hard and", "and",
    "achieve", "puzzle", "moonlight", "joy", "star", "courage", "greatness!", "dreams", "will", "bloom"
]

def extract_encouragement(word_list):
    pass

```

```python
grader.check("q")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q (points: None)
```
>>> words = ['banana', 'table', 'you', 'are', 'cloud', 'doing', 'river', 'great', 'apple', 'work', 'keep', 'on', 'shining', 'this', 'and', 'trust', 'in', 'assignment;', 'path', 'keep', 'always', 'believe', 'mirror', 'working', 'can', 'prosper', 'anything', 'with', 'kindness', 'hard and', 'and', 'achieve', 'puzzle', 'moonlight', 'joy', 'star', 'courage', 'greatness!', 'dreams', 'will', 'bloom']
>>> assert extract_encouragement(words) == 'you are doing great on this assignment; keep working hard and achieve greatness!'
>>> assert words == ['you', 'are', 'doing', 'great', 'on', 'this', 'assignment;', 'keep', 'working', 'hard and', 'achieve', 'greatness!']
```
```
>>> words2 = ['a', 'b', 'c', 'd', 'e', 'f', 'g']
>>> assert extract_encouragement(words2) == 'c d f'
>>> assert words2 == ['c', 'd', 'f']
>>> assert extract_encouragement(words2) == 'f'
>>> assert words2 == ['f']
>>> assert extract_encouragement(words2) == ''
>>> assert words2 == []
```


Tests for q1 (points: 2)
```
>>> assert return_div_3([1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 12, 15, 21]) == [3, 6, 9, 12, 15, 21]
```
```
>>> assert return_div_3(['hello', 12, 15, 'yes', 21]) == [12, 15, 21]
```
